# lab_simulation: 1/14,000,605

As you learned in the lecture, **simulation** is an extremely powerful tool to estimate probabilities by replicating the reality and determining the successful observations.

In this lab, we will try some increasingly interesting simulations and explore the possibilities with this tool.

(Wonder why the title is 1/14,000,605? See Avengers: Infinity War at [https://www.youtube.com/watch?v=eGKPfZTXHsc](https://www.youtube.com/watch?v=eGKPfZTXHsc))

# Part 0 | Your Group!

Edit the next Python cell to add information about who you're working within your lab section:

```python
# First, meet your CAs and TA if you haven't already!
# ...first name is enough, we'll know who they are! :)
ta_name = "Esther"
ca1_name = "Ananya"
ca2_name = "Leon"


# Also, make sure to meet your team for this lab! Find out their name, what major they're in,
# and learn something new about them that you never knew before!
partner1_name = "Luke"
partner1_netid = "mateski2"
partner1_major = "Econometrics"

partner2_name = "Amal"
partner2_netid = "akhati5"
partner2_major = "HDFS"

partner3_name = ""
partner3_netid = ""
partner3_major = ""
```

## Example 0.1: Random seed (IMPORTANT: PLEASE READ AND RUN THIS EXAMPLE)

A random seed is a number used to initialize a pseudorandom number generator. If we use the same seed, we will get the same random number sequence.

In this lab, we will use the random seed to help you check the simulations' output. Please do not delete the line that set the seed to 107 in any puzzles.

```python
import pandas as pd
import random as rd
rd.seed(107)
print(rd.random())
print(rd.random())
print(rd.random())
rd.seed(107)
print(rd.random())
print(rd.random())
```

Output:
```
0.24648195966935815
0.5812899303907797
0.9722810069256511
0.24648195966935815
0.5812899303907797
```

## Example 0.2: Simulation of two dice (you can go to part 1 directly if you are confident)

Suppose we roll two different **fair six-sided dice** at the same time.  One die is colored **blue** while the other is colored **red**.

Simulate the above situation **1,000** times and store the observations of the value of the blue die and the red die in `df0`.

<!-- grade_id: cell-c93866d5164aa8aa, points: 0, grade: false, solution: true -->
```python
# Step 0: Set the random seed for replicability
rd.seed(107)
# Step 1: Start with an empty list to store the simulated data
data = []
# Step 2: Implement the simulation with a for-loop
for i in range(1000):
    blue = rd.randint(1, 6)
    red = rd.randint(1, 6)
    d = {"blue": blue, "red": red}
    data.append(d)    
# Step 3: Convert the list of simulated data to a dataframe
df0 = pd.DataFrame(data)
# Sanity check: Show a few random rows
df0.sample(5)
```

Output:
```
blue  red
710     6    1
200     5    1
979     3    4
998     5    6
738     2    1
```

## Example 0.3: Probability estimation of a two-dice event

Using the above simulated data, we can estimate the probability that **the blue die lands on 4 or the red die lands on 2**.

<!-- grade_id: cell-5144eadebe059446, points: 0, grade: false, solution: true -->
```python
# Numerator:    sum only counts the number of TRUEs.
# Denominator:  len counts both TRUEs and FALSEs.
#               You can try len((df0['blue'] == 4) | (df0['red'] == 2)).
p03 = sum((df0['blue'] == 4) | (df0['red'] == 2))/len(df0)
print(p03)
# An alternative approach is to subset the success dataframe and find its length.
df0_success = df0[(df0['blue'] == 4) | (df0['red'] == 2)]
p03_alt = len(df0_success)/len(df0)
print(p03_alt)
```

Output:
```
0.318
0.318
```

## Example 0.4: Exact probability of a two-dice event

We can also calculate the exact probability that **the blue die lands on 4 or the red die lands on 2**, which is not far from our estimate in Example 0.3.

<!-- grade_id: cell-a28ce929e8600774, points: 0, grade: false, solution: true -->
```python
p03_exact = 1/6 + 1/6 - 1/36
p03_exact
```

Output:
```
0.3055555555555555
```

# Part 1 | Rolling Three Dice

## Puzzle 1.1: Simulation of three dice

Suppose we roll three different **fair six-sided dice**: a **white**, a **red**, and a **blue**.

Simulate the above situation **1,000** times and store the observations in `df1`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">see Example 0.2. Try not to copy and paste so that you can learn line by line.</span>

<!-- grade_id: cell-c42efbf1ba47dcab, points: 0, grade: false, solution: true -->
```python
rd.seed(107)
data = []
for i in range (1000):
    white = rd.randint(1, 6)
    red = rd.randint(1, 6)
    blue = rd.randint(1, 6)
    d = {"white": white, "red": red, "blue": blue}
    data.append(d) 
df1 = pd.DataFrame(data)
df1.sample(5)
```

Output:
```
white  red  blue
990      3    5     1
565      4    4     3
292      3    5     5
66       4    5     3
351      5    5     6
```

## Puzzle 1.2: Probability estimation of a three-dice event

Estimate the probability that the **sum of all three dice is exactly 9**. Store the result in `p12`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">see Example 0.3. You can use .sum(axis=1) to obtain the sum of all three dice row by row.</span>

<!-- grade_id: cell-b16fd8d9c98834ef, points: 0, grade: false, solution: true -->
```python
sumOfDice = df1.sum(axis=1)
p12 = sum(sumOfDice==9) / len(df1) 
p12
```

Output:
```
0.127
```

<!-- grade_id: cell-571231fcdfa9de08, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 1 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any error, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell(s), make changes, and RE-RUN your code and then this cell.

assert(len(df1) == 1000), "Make sure your df1 has exactly 1,000 observations"
assert(p12 == 127/1000), "Make sure you set the seed to 107 and estimate the correct probability in p12"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

# Part 2 | Flipping Four Coins

## Puzzle 2.1: Simulation of four coins

Suppose we flip **four fair coins**, one coin at a time and one after another. Each coin has two sides: **"head" (defined as 1)** and **"tail" (defined as 0)**.

Simulate the above situation **50,000** times and store the observations in `df2`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">see Example 0.2. Try not to copy and paste so that you can learn line by line.</span>

<!-- grade_id: cell-da32a7ecb4863a7f, points: 0, grade: false, solution: true -->
```python
rd.seed(107)
data = []
for i in range (50000):
    coin1 = rd.randint(0, 1)
    coin2 = rd.randint(0, 1)
    coin3 = rd.randint(0, 1)
    coin4 = rd.randint(0, 1)
    d = {"coin1": coin1, "coin2": coin2, "coin3": coin3, "coin4": coin4}   
    data.append(d)
df2 = pd.DataFrame(data)
df2.sample(5)
```

Output:
```
coin1  coin2  coin3  coin4
40014      1      1      0      0
24343      0      1      0      0
29476      0      1      1      1
41098      1      0      1      0
34469      1      1      0      1
```

## Puzzle 2.2: Probability estimation of a four-coins event

Estimate the probability that **the first two coin flips are both heads and last two coin flips are both tails**. Store the result in `p22`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">see Example 0.3.</span>

<!-- grade_id: cell-500cf95087dcb6db, points: 0, grade: false, solution: true -->
```python
p22 = sum((df2.coin1 == 1) & (df2.coin2 == 1) & (df2.coin3 == 0) & (df2.coin4 == 0)) / len(df2)
p22
```

Output:
```
0.06382
```

## Puzzle 2.3: Probability estimation of another four-coins event

Estimate the probability that **at most two coin flips are heads**. Store the result in `p23`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">see Example 0.3. You can use .sum(axis=1) to obtain the "sum" of four coins row by row.</span>

<!-- grade_id: cell-500cf95087dcb6db, points: 0, grade: false, solution: true -->
```python
num = df2.sum(axis=1)
p23 = sum((num==2) | (num==1) | (num==0)) / len(df2)
p23
```

Output:
```
0.68536
```

<!-- grade_id: cell-ee8226d0066bb603, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 2 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any error, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell(s), make changes, and RE-RUN your code and then this cell.

assert(len(df2) == 50000), "Make sure your df2 has exactly 50,000 observations"
assert(p22 == 3191/50000), "Make sure you set the seed to 107 and estimate the correct probability in p22"
assert(p23 == 34268/50000), "Make sure you set the seed to 107 and estimate the correct probability in p23"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

# Part 3 | Avengers: Infinity War

## Puzzle 3.1: Simulation of the battle on Titan

Suppose the **power level** of the characters engaged in the battle on Titan can be modelled as follows:

1. Thanos:  a random integer between 80 to 120.
2. Strange: a random integer between 10 to 20.
3. Stark:   a random integer between 10 to 16.
4. Parker:  a random integer between 6 to 12.
5. Mantis:  a random real number between 0.8 to 1 (inclusive). The code for Mantis is given.
6. Drax:    a random integer between 8 to 16.
7. Nebula:  a random integer between 8 to 16.
8. Quill:   a random integer between -10 to 10.

Simulate the power levels **100,000** times **in the order listed above** and store the observations in `df3`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">see Example 0.2. Try not to copy and paste so that you can learn line by line.</span>

<!-- grade_id: cell-da32a7ecb4863a7f, points: 0, grade: false, solution: true -->
```python
rd.seed(107)
data = []
for i in range(100000):
    thanos = rd.randint(80, 120)
    strange = rd.randint(10, 20)
    stark = rd.randint(10, 16)
    parker = rd.randint(6, 12)
    mantis = rd.uniform(0.8, 1)
    drax = rd.randint(8, 16)
    nebula = rd.randint(8, 16)
    quill = rd.randint(-10, 10)
    d = {"thanos": thanos, "strange": strange, "stark": stark, "parker": parker, "mantis": mantis, "drax": drax, "nebula": nebula, "quill": quill}
    data.append(d)
df3 = pd.DataFrame(data)
df3.sample(5)
```

Output:
```
thanos  strange  stark  parker    mantis  drax  nebula  quill
70516      90       13     13      11  0.996017    11      11     -3
44746      85       18     15      12  0.979204     8      14     -7
83090     109       15     15      12  0.938733    16      12      1
26958     105       19     13      10  0.889213    14       9     -7
38463     116       10     15      11  0.859083     9       8     10
```

## Puzzle 3.2: Probability estimation of winning the battle

Suppose the superheroes win when
$$
\mathrm{Strange} +\mathrm{Stark} +\mathrm{Parker} +\mathrm{Drax} +\mathrm{Nebula} +\mathrm{Quill} \ge \mathrm{Mantis} \times \mathrm{Thanos}.
$$
Estimate the probability that **the superheroes win**. Store the result in `p32`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">see Example 0.3. You can use df[list].sum(axis=1) to obtain the row sum of columns with name in list.</span>

<!-- grade_id: cell-500cf95087dcb6db, points: 0, grade: false, solution: true -->
```python
heroes = ["strange", "stark", "parker", "drax", "nebula", "quill"]
df3["superhero power"] = df3[heroes].sum(axis=1)
df3["villain power"] = df3["mantis"] * df3["thanos"]


p32 = sum(df3["superhero power"] >= df3["villain power"]) / len(df3)
p32
```

Output:
```
0.01717
```

## Puzzle 3.3: Conditional probability estimation of winning the battle

Recall the superheroes win when
$$
\mathrm{Strange} +\mathrm{Stark} +\mathrm{Parker} +\mathrm{Drax} +\mathrm{Nebula} +\mathrm{Quill} \ge \mathrm{Mantis} \times \mathrm{Thanos}.
$$
Estimate the probability that **the superheroes win given that $\mathrm{Mantis} < 0.9$**. Store the result in `p33`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">use Bayes' rule (or definition of conditional probability) on top of the code in Puzzle 3.2.</span>

<!-- grade_id: cell-500cf95087dcb6db, points: 0, grade: false, solution: true -->
```python
df_mantis = df3[df3["mantis"] < 0.9]

p33 = sum(df_mantis["superhero power"] >= df_mantis["villain power"]) / len(df_mantis)
p33
```

Output:
```
0.03127697588982796
```

<!-- grade_id: cell-ee8226d0066bb603, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 3 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any error, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell(s), make changes, and RE-RUN your code and then this cell.

assert(len(df3) == 100000), "Make sure your df3 has exactly 100,000 observations"
assert(p32 == 1717/100000), "Make sure you follow the instruction in Puzzle 3.1 and estimate the correct probability in p32"
assert(p33 == 1558/49813), "Make sure you follow the instruction in Puzzle 3.1 and estimate the correct probability in p33"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

## Puzzle 3.4: Usefulness of simulation

❓ **Individual Reflection Question** ❓ Can you find the exact probabilities in Puzzle 3.2/3.3? What are the advantage(s) of simulation?

It would be hard to find the exact probabilities. Using math equations, you would need to find the probability of all kinds of different combinations. With simulation, it is much easier because the computer can randomly create combinations. Although the simulation will likely not give the exact probability due to sampling variance, the results can get very close to the exact probability with more samples. Thus, the advantage of simulation is that they can test several trials to estimate probabilities of events that involve large amounts of variables.

❓ **Group Discussion Question** ❓ Recall our title 1/14,000,605 from Avengers: Infinity War (see https://www.youtube.com/watch?v=eGKPfZTXHsc if you haven't). Discuss with your groupmates if you have seen other examples of simulation in movies.

# Part 4 | Beyond probability estimation

## Puzzle 4.1: Simulation of points on a unit square

Imagine we have a **unit square** with the corners at $(0,0)$, $(0,1)$, $(1,1)$ and $(1,0)$, respectively.

Simulate **100,000** random points on this square and store the observations in `df4`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">see Example 0.2. The x-coordinate and y-coordiante of a random point can be generated by rd.random().</span>

<!-- grade_id: cell-da32a7ecb4863a7f, points: 0, grade: false, solution: true -->
```python
rd.seed(107)
data = []
for i in range (100000):
    x = rd.random()
    y = rd.random()
    d = {"x": x, "y": y}
    data.append(d)
df4 = pd.DataFrame(data)
df4.sample(5)
```

Output:
```
x         y
88303  0.787140  0.849252
92382  0.076881  0.779602
12675  0.086704  0.312691
2404   0.230919  0.590708
52139  0.232231  0.230774
```

## Puzzle 4.2: Estimation of $\pi$

Now, we incribe a **circle** (with radius $0.5$) in the previous square. The circumcenter is $(0.5, 0.5)$. Clearly, a point would fall inside the circle if
$$
(x-0.5)^2 +(y-0.5)^2 \le 0.5^2.
$$
On the other hand, we have the following relationship:
$$
\frac{\text{no. of pts inside the circle}}{\text{no. of pts inside the square}}
\approx \frac{\text{area of the circle}}{\text{area of the square}}
= \frac{0.5^2 \pi}{1} = \frac{\pi}{4}.
$$
Rearranging the equation, we can estimate $\pi$ by
$$
\pi \approx \frac{4 \times \text{no. of pts inside the circle}}{\text{no. of pts inside the square}}.
$$

Estimate $\pi$ using the above relationship. Store the result in `pie`.

Hint (double-click this cell to see it): <span style="color:#ffffff00">see Example 0.3. 
The number of points inside the circle can be found by checking the equation (x-0.5)^2 +(y-0.5)^2 <= 0.5^2.</span>

<!-- grade_id: cell-500cf95087dcb6db, points: 0, grade: false, solution: true -->
```python
pie = 4*(sum(((df4["x"] - 0.5)**2 + (df4["y"] - 0.5)**2) <= (0.5**2)) / len(df4))
pie
```

Output:
```
3.14192
```

<!-- grade_id: cell-ee8226d0066bb603, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Part 4 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any error, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell(s), make changes, and RE-RUN your code and then this cell.

assert(len(df4) == 100000), "Make sure your df4 has exactly 100,000 observations"
assert(pie == 314192/100000), "Make sure you use the correct equation(s) in Puzzle 4.2"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

## Puzzle 4.3: Beyond probability estimation

❓ **Individual Reflection Question** ❓ Puzzle 4.2 presents an example where we use simulation for estimating quantity other than probability. Can you name other application(s) of simulation?

Hint (double-click this cell to see it): <span style="color:#ffffff00">check https://en.wikipedia.org/wiki/Monte_Carlo_method#Applications if you have no clues at all.</span>

We use simulation a lot in chemistry. For example, the PHET website has a gas properties simulator. It lets us re-enact real-life situations that controls the volume/temperature/amount of gas in a box. We can see the gas pressure change depending on the environment conditions.

We could also apply simulation to the lottery. We could simulate someone buying a certain number of tickets (let's say 10 tickets). Then we set certain tickets to be the winners, making sure that only a certain # of tickets are winners (ex. if the odds of winning is 1/10, we can set 100 out of 1000 ticket numbers to be winners). We can then see how many times they get winning tickets in the lottery.

❓ **Group Discussion Question** ❓ Do you know any other approaches to estimate $\pi$? Discuss with your groupmates whether you like the simulation approach or not (and why).

# Finale!

You're almost done -- congratulations!

You need to do two more things:

1.  Save your work. To do this, go to File -> Save All

2.  After you have saved, exit this notebook and follow the webpage instructions to commit this lab to your Git repository!